# Assignment 04: Taylor series and Newton Raphson
The purpose of this assignment is to develop your skills in creating approximate functions using Taylor series and find roots of functions using the Newton's iterations.
 

Taylor series expansion of 
$ f(x) = sin(ax) $ near 0, where a is a known constant:

$ax - \frac{a^3}{3!}x^3+ \frac{a^5}{5!}x^5 - \frac{a^7}{7!}x^7 + \cdots$ or
$ \sum_{n = 0}^{\infty}\frac{a^{2n+1}(-1)^n}{(2n+1)!}x^{2n+1} $

## Q1
Write the Taylor's series expansion of the function $f(x)
	= sin(ax)$ 
	near $x = 0$,
	where $a \ne 0$ is a known constant. Write a Python function to compute the approximate solution and error. Compute the relative and absolute errors using $math.sin(ax)$ as the exact answer.
	

```python
import math
def taylor_sin_ax(x, a, n): #n define the number of terms of Taylor series used
    ...
    return approx, abs_error, rel_error
    
x = math.pi/4
a = 2
print(taylor_sin_ax(x, a, 3))
print(taylor_sin_ax(x, a, 5))
```

## Q2
Determine the root of $f(x) = x - 2e^{-x}$ by:<br>
### (a)
Using the bisection method. Start with $a = 0$ and $b = 1$, and compute **the root and the error at the end of the first three iterations.**<br>
### (b)
Using the Newton's method, Start at $x_0 = 1$ and compute **the root and the error at the end of the the first three iterations**. Compare the error against the bisection approach.

```python
# (a)
# store the root and error in the variable 'root1' and 'error1'
print(root1, error1)
```

```python
# (b)
# store the root and error in the variable 'root2' and 'error2'
print(root2, error2)
```

## Q3
Using the Newton-Raphson iteration find the **root** of the equation $\sqrt{x} + x^2 = 7$. With an initial guess of $x = 7$ compute the **number of iterations** required for the error in the root to be below $1.0e^{-6}$.


```python
tol = 1e-6

print(root3, num_iter)
```

Tests for q1 (points: 4)
```
>>> import math
>>> x = math.pi / 4
>>> a = 2
>>> assert round(taylor_sin_ax(x, a, 3)[0] - 1.0045248555348174, 9) == 0
>>> assert round(taylor_sin_ax(x, a, 3)[1] - 0.004524855534817407, 9) == 0
>>> assert round(taylor_sin_ax(x, a, 3)[2] - 0.004524855534817407, 9) == 0
>>> assert round(taylor_sin_ax(x, a, 5)[0] - 1.0000035425842861, 9) == 0
>>> assert round(taylor_sin_ax(x, a, 5)[1] - 3.542584286142514e-06, 9) == 0
>>> assert round(taylor_sin_ax(x, a, 5)[2] - 3.542584286142514e-06, 9) == 0
```
```
>>> import random
>>> n = 100
>>> a = 2
>>> for i in range(20):
...     x = random.random()
...     assert round(taylor_sin_ax(x, a, n)[0] - math.sin(a * x), 9) == 0
```


Tests for q2 (points: 3)
```
>>> assert root1 == 0.875
>>> assert round(error1 - 0.04127596064298322, 9) == 0
```
```
>>> assert round(root2 - 0.8526055020070299, 9) == 0
>>> assert round(error2 - 6.695596982741675e-12, 15) == 0
```


Tests for q3 (points: 2)
```
>>> assert round(root3 - 2.338940066881352, 9) == 0
>>> assert num_iter == 5
```
