# Assignment 05: Vectors and Matrices
The purpose of this assignment is to develop your skills in using Numpy arrays and doing vector and matrix operations.

## Q1
Write custom individual functions to perform the following matrix operations, without using the in-built numpy functions (except np.shape function). Add checks on suitability of the shape of matrices whenever possible <br>

<li>Add two matrices</li>
<li>Subtract two matrices</li>
<li>Scalar product of a matrix</li>
<li>Multiply two matrices</li>
<li>Sum of all elements in a matrix</li>
<li>Check if a given matrix is symmetric and square.</li>

```python
import numpy as np
# Add two matricies
def matrix_add(A, B):
    ...
# Subtract two matrices
def matrix_subtract(A, B):
    ...
# Scalar product of a matrix
def matrix_scalar_product(c, A):
    ...
# Multiply two matrices
def matrix_product(A, B):
    ...
# Sum of all elements in a matrix
def matrix_sum_elements(A):
    ...
# Check if a given matrix is symmetric and square 
def symm_sq_check(A):
    ...
```

```python
# here are some input samples to test your functions. 
A = np.array([[1, 2], [3, 4]])
B = np.array([[5, 6], [6, 8]])
```

## Q2
Let $A$ be a 4 x 4 matrix and B a 2 x 1 matrix. The symbol $\cdot$ represents a dot product. Using the shape of the matrix as a guide, indicate if the following expressions are valid or not.<br>
1. $ A + B$
2. $ A \cdot B$
3. $ A \cdot A$
4. $ B \cdot B$
5. $ B^T \cdot B \cdot A$

```python
q2_ans = [...] #store boolean values to represent the validity of each expression.
```

## Q3
Using the following matrices and vectors. Compute the following.
\begin{align*}
	A & =  \begin{bmatrix}
		1 & 4 & -2 \\
		4 & 8 & 6 \\
		-2 & 6 & 12 
    	\end{bmatrix} \\
	%
	B & =  \begin{bmatrix}
	6 & 2 & -2 \\
	4 & 8 & 3 \\
	-1 & 6 & 9 
	\end{bmatrix} \\
	c & = \left[6, -4, 3\right]\\
	d & = \left[3, -1, 5\right]
	\end{align*}
    
1. $ A - A^T$
2. $A \cdot B$
3. $8A - 4B$
4. $B^T \cdot B$
5. $c^T \cdot B$
6. $A\cdot B \cdot c \cdot d^T$
7. $c \times d$
8. $A^{-1}$ 
9. check if $A * A^{-1} = I$

```python
A = np.array([[1, 4, -2], [4, 8, 6], [-2, 6, 12]])
B = np.array([[6, 2, -2], [4, 8, 3], [-1, 6, 9]])
c = np.array([6, -4, 3])
d = np.array([3, -1, 5])
q3_ans = [] # store answers in this list

#hint: use np.linalg.inv() to get inverse of A
#hint: use np.allclose() and its default tolerance to check if A*A^-1 equal to I
```

## Q4
Find a a unit vector (i.e., vector of magnitude equal to 1) that is perpendicular to both $c$ and $d$. Hint: *use the cross product definition*.


```python
unit_vector = ...
```

## Q5
Using array slicing on matrix A
\begin{align*}
	A = \begin{bmatrix}
4.0 & 7.0 & -2.43 & 67.1 \\
-4.0 & 64.0 & 54.7 & -3.33 \\
2.43 & 23.2 & 3.64 & 4.11 \\
1.2 & 2.5 & -113.2 & 323.22
	\end{bmatrix}
\end{align*}
    
<li>Extract the third column as a 1D array
</li><li>Extract the first two rows as a 2D sub-array
</li><li>Extract the bottom-right $2 \times 2$ block as a 2D sub-array
</li><li>Sum the last column</li>

```python
A = np.array([[4.0, 7.0, -2.43, 67.1], [-4.0, 64.0, 54.7, -3.33],
              [2.43, 23.2, 3.64, 4.11], [1.2, 2.5, -113.2, 323.22]])
q5_ans = [] #store answers in this list
```

Tests for q1 (points: 2)
```
>>> A = np.array([[1, 2], [3, 4]])
>>> B = np.array([[5, 6], [6, 8]])
>>> assert np.array_equal(matrix_add(A, B), A + B)
>>> assert np.array_equal(matrix_subtract(A, B), A - B)
>>> assert np.array_equal(matrix_scalar_product(2, A), 2 * A)
>>> assert np.array_equal(matrix_product(A, B), np.dot(A, B))
>>> assert round(matrix_sum_elements(A) - np.sum(A), 5) == 0
>>> assert symm_sq_check(A) == np.array_equal(A.transpose(), A)
>>> assert symm_sq_check(B) == np.array_equal(B.transpose(), B)
```
```
>>> for i in range(20):
...     A = np.random.random((5, 5))
...     B = np.random.random((5, 5))
...     assert np.array_equal(matrix_add(A, B), A + B)
...     assert np.array_equal(matrix_subtract(A, B), A - B)
...     assert np.array_equal(matrix_scalar_product(2, A), 2 * A)
...     assert np.array_equal(matrix_product(A, B).round(4), np.dot(A, B).round(4))
...     assert round(matrix_sum_elements(A) - np.sum(A), 9) == 0
```


Tests for q2 (points: 2)
```
>>> q2 = [False, False, True, False, False]
>>> for i in range(len(q2)):
...     assert q2[i] == q2_ans[i], f'Answer {i + 1} is incorrect'
```


Tests for q3 (points: 2)
```
>>> q3 = [np.array([[0, 0, 0], [0, 0, 0], [0, 0, 0]]), np.array([[24, 22, -8], [50, 108, 70], [0, 116, 130]]), np.array([[-16, 24, -8], [16, 32, 36], [-12, 24, 60]]), np.array([[53, 38, -9], [38, 104, 74], [-9, 74, 94]]), np.array([17, -2, 3]), -352, np.array([-17, -21, 6]), np.array([[-0.23076923, 0.23076923, -0.15384615], [0.23076923, -0.03076923, 0.05384615], [-0.15384615, 0.05384615, 0.03076923]]), True]
>>> for i in range(len(q3) - 2):
...     assert np.array_equal(q3_ans[i], q3[i]), f'Answer {i + 1} is incorrect'
>>> assert np.array_equal(q3_ans[-2].round(5), q3[-2].round(5)), f'Answer 8 is incorrect'
>>> assert q3_ans[-1] == q3[-1], f'Answer 9 is incorrect'
```


Tests for q4 (points: 2)
```
>>> np.allclose(unit_vector, [-0.61423497, -0.75876084, 0.21678881])
True
```


Tests for q5 (points: 2)
```
>>> q5 = [np.array([-2.43, 54.7, 3.64, -113.2]), np.array([[4.0, 7.0, -2.43, 67.1], [-4.0, 64.0, 54.7, -3.33]]), np.array([[3.64, 4.11], [-113.2, 323.22]]), 391.1]
>>> for i in range(len(q5)):
...     assert np.array_equal(q5[i], q5_ans[i]), f'Answer {i + 1} is incorrect'
```
