## Exercise 02 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - If $f(x_0) \cdot f(x_{\rm mid}) < 0$: 

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between 
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - Else

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = x^3 - 6x^2 + 4x + 12
$$

has one root somewhere between $x_0 = 3$ and $x_1 = 6$.

1. Use the bisection method to find an approximate root $x_{r}$ using 15 iterations 
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that 
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations 
   required (use a `for` loop with a large number of iterations).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`. And store them in the list `ans`= [`x_mid`, `f`]. For task 2, you will store the number of iteration in `n_iter`

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`. 

```python
# Task 1: For loop with a defined number of iterations
# Initial end points
x0 = 3.0
x1 = 6.0

for n in range(...):
    ...
ans1 = [x_mid, f] # store answers in the list
# Print the number of iterations, final x and final f value
print(n+1, ans1[0], ans1[1])
```

```python
assert round(ans1[0] - 4.534149169921875, 10) == 0.0
assert abs(ans1[1]) < 0.0009
```

```python
# Task 2: For loop with a defined tolerance
# Initial end points
x0 = 3.0
x1 = 6.0
tol = 1.0e-6

n_iter = ... #number of iteration
ans2 = [x_mid, f] # store answers in the list
print(n_iter, ans2[0], ans2[1]) 

```

```python
assert n_iter == 23
assert round(ans2[0] - 4.534149169921875, 3) == 0.0
assert abs(ans2[1]) < 1.0e-6
```

## Exercise 02 (series expansion)

The power series expansion for the sine function is: 

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `for` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$. 

Store the absolute value of the error in the variable `error_1` and `error_2`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use 
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the absolute error.

```python
# Import the math module to access math.sin and math.factorial
import math

# Value at which to approximate sine
x = 1.5*math.pi

# Initialise approximation of sine
approx_sin = 0.0

# YOUR CODE HERE
...
    
print("The error is:")
print(error_1)
```

```python
assert error_1 < 1.0e-12
```

```python
# Import the math module to access math.sin and math.factorial
import math

# Value at which to approximate sine
x = 1.5*math.pi

# Tolerance and initial error (this just needs to be larger than tol)
tol = 1.0e-8
error_2 = tol + 1.0

# Intialise approximation of sine
approx_sin = 0.0

# Initialise counter
n_term = 0

# Loop until error satisfies tolerance or a fixed number of iterations is reached
for i in range(1001):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
print("The error is:")
print(error_2)

print("Number of terms in series:")
print(n_term)
```

```python
assert error_2 <= 1.0e-8
assert n_term == 12
```

Tests for q1 (points: 5)
```
>>> assert round(ans1[0] - 4.534149169921875, 10) == 0.0
```
```
>>> assert abs(ans1[1]) < 0.0009
```
```
>>> assert n_iter == 23
```
```
>>> assert round(ans2[0] - 4.534149169921875, 3) == 0.0
```
```
>>> assert abs(ans2[1]) < 1e-06
```


Tests for q2 (points: 5)
```
>>> assert error_1 < 1e-12
```
```
>>> assert error_2 <= 1e-08
```
```
>>> assert n_term == 12
```
