## Exercise 04.1 (selecting and passing data structures)

Given the coordinates of the vertices of a triangle, $(x_0, y_0)$, $(x_1, y_1)$ and $(x_2, y_2)$, the area $A$ of the triangle is given by:
$$
A = \left| \frac{x_0(y_1  - y_2) + x_1(y_2 - y_0) + x_2(y_0 - y_1)}{2} \right|
$$
There is an old function named `area` that computes the area of a triangle given the coordinates of the vertices as six arguments ($x$ and $y$ components of each vertex). The order of the function arguments must be (`x0, y0, x1, y1, x2, y2)`. With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

You have been hired as a Python consultant to rewrite the function for computing the area of a triangle.

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.<br>
**HINT:** <li>if the number of input coordinates is not 3, print `incorrect number of coordinates`</li><li>if the calculated area is 0, print `invalid coordinates for a triangle` </li>

```python
# BEGIN SOLUTION NO PROMPT
def area_triangle(coord):
  """ calculates the area of a triangle """

  if len(coord) != 3:
    print("incorrect number of coordinates")
    return

  x0 = coord[0][0]
  y0 = coord[0][1]
  x1 = coord[1][0]
  y1 = coord[1][1]
  x2,y2 = coord[2]

  area = abs((x0*(y1-y2)+x1*(y2-y0)+x2*(y0-y1))/2)

  if area == 0:
    print("invalid coordinates for a triangle")
    return
  
  return area
# END SOLUTION
""" # BEGIN PROMPT
# input format
# coord = [(x0,y0),(x1,y1),(x2,y2)]
# area_triangle(coord)
# YOUR CODE HERE
def area_triangle(...):
    ...
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\n# input format\n# coord = [(x0,y0),(x1,y1),(x2,y2)]\n# area_triangle(coord)\n# YOUR CODE HERE\ndef area_triangle(...):\n    ...\n'
```

```python
a = [(1,1),(2,5),(3,1)]
area_triangle(a) # should return 4.0
```

Output:
```
4.0
```

```python
a = [(1,1),(2,5)]
area_triangle(a) # should print "incorrect number of coordinates"
```

Output:
```
incorrect number of coordinates
```

```python
a = [(1,1),(2,5),(3,1),(2,0)]
area_triangle(a) # should print "incorrect number of coordinates"
```

Output:
```
incorrect number of coordinates
```

```python
a = [(1,1),(2,5),(2,5)]
area_triangle(a) # should print "invalid coordinates for a triangle"
```

Output:
```
invalid coordinates for a triangle
```

## Exercise 04.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

```python
# BEGIN SOLUTION NO PROMPT
def polygon_area(vertices):
  """Compute the area of a polygon"""
  
  if (len(vertices) < 3):
    print("Invalid number of coordinates to compute an area")
    return

  area = 0
  # approach 1
  for i in range(len(vertices)):
    xi = vertices[i][0]
    yi = vertices[i][1]
    if i == len(vertices)-1:
      xi_1 = vertices[0][0]
      yi_1 = vertices[0][1]
    else:
      xi_1 = vertices[i+1][0]
      yi_1 = vertices[i+1][1]
    
    area += 0.5*(xi*yi_1-xi_1*yi)
  # approach 2
  # vertices.append(vertices[0])  # append the first vertex at the end
  # for i in range(len(vertices)-1):
  #   xi = vertices[i][0]
  #   yi = vertices[i][1]
  #   xi_1 = vertices[i+1][0]
  #   yi_1 = vertices[i+1][1]
  #   area += 0.5*(xi*yi_1-xi_1*yi)
    
  return abs(area)
# END SOLUTION
""" # BEGIN PROMPT
def polygon_area
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\ndef polygon_area\n'
```

```python
vertices = [(1,1),(2,5),(3,1)]
polygon_area(vertices) # should return 4.0
```

Output:
```
4.0
```

```python
vertices=[(1,1),(2,5),(5,5),(6,1)]
polygon_area(vertices) # should return 16.0
```

Output:
```
16.0
```

## Exercise 04.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

```python
# BEGIN SOLUTION NO PROMPT
def sum_vector(x, y):
  """Return sum of two vectors"""
  if (len(x) != len(y)):
    print("Invalid dimensions of x and y. They should match")
    return
  
  z = [0]*len(x)
  for i in range(len(x)):
    z[i] = x[i] + y[i]
  
  return z
# END SOLUTION
""" # BEGIN PROMPT
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
def sum_vector(...):
    ...
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\na = [0, 4.3, -5, 7]\nb = [-2, 7, -15, 1]\ndef sum_vector(...):\n    ...\n'
```

```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
sum_vector(a,b) # should return [-2, 11.3, -20, 8]
```

Output:
```
[-2, 11.3, -20, 8]
```

## Exercise 04.4 (dictionaries)

Create a dictionary that maps build names (the key) to building abbreviations for at least 5 buildings
(you can find abbreviations at https://utdirect.utexas.edu/apps/campus/buildings/nlogon/facilities/UTM/).
From the dictionary, produce and print

1. A dictionary from building abbreviation to name; and
1. A list of building abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps building names (the key) to dictionaries of:

- Building abbreviation
- Year occupied
- Zones
 
for at least 5 building. Take the data from https://utdirect.utexas.edu/apps/campus/buildings/nlogon/facilities/UTM/. Using this dictionary, 

1. Find the oldest occupied building, and print the zone and the abbreviation for this building
2. Find the oldest building in a given zone.

```python
# BEGIN SOLUTION NO PROMPT
buildings = ["ATHLETIC FIELDS PAVILION (EASTSIDE)", 
            "ALMETRIS DUREN RESIDENCE HALL",
            "ATHLETIC FIELDS PAVILION (REHAB)",
            "ATHLETIC FIELDS PAVILION",
            "ANNA HISS GYMNASIUM"]
building_abv = {"ATHLETIC FIELDS PAVILION (EASTSIDE)":"AF2",
             "ALMETRIS DUREN RESIDENCE HALL":"ADH" ,
             "ATHLETIC FIELDS PAVILION (REHAB)":"AF1",
             "ATHLETIC FIELDS PAVILION":"AFP",
             "ANNA HISS GYMNASIUM":"AHG"}
abv_building = {}
abv_building = {value:key for key, value in building_abv.items()}

building_year = {"ALMETRIS DUREN RESIDENCE HALL":        2007,
                  "ATHLETIC FIELDS PAVILION (REHAB)":    2012,
                  "ATHLETIC FIELDS PAVILION (EASTSIDE)": 2012,
                  "ANNA HISS GYMNASIUM":                 1930,
                  "ATHLETIC FIELDS PAVILION":            1995}

building_zone = {"ALMETRIS DUREN RESIDENCE HALL":        2,
                  "ATHLETIC FIELDS PAVILION (REHAB)":    3,
                  "ATHLETIC FIELDS PAVILION (EASTSIDE)": 3,
                  "ANNA HISS GYMNASIUM":                 1,
                  "ATHLETIC FIELDS PAVILION":            3}
sorted_abv = sorted(building_abv.values())
oldest_year = 2050
for building, year in building_year.items():
    if year < oldest_year:
        old_building = building
        oldest_year = year
old_zone = building_zone[old_building]
old_abv = building_abv[old_building]

oldest_year = 2050
for building, year in building_year.items():
    if building_zone[building] == 3:
        if year < oldest_year:
            old_building2 = building
            oldest_year = year
# END SOLUTION
""" # BEGIN PROMPT
# find the abbreviations for the buildings
buildings = ["ATHLETIC FIELDS PAVILION (EASTSIDE)", 
            "ALMETRIS DUREN RESIDENCE HALL",
            "ATHLETIC FIELDS PAVILION (REHAB)",
            "ATHLETIC FIELDS PAVILION",
            "ANNA HISS GYMNASIUM"]
# make a dictionary that maps building names (key) to building abbreviations (value)
building_abv = ...

# make a dictionary that maps building abbreviations (key) to building names (value)
# Hint: use building_abv to create the this dictionary. Don't manually type keys and values...
abv_building = ...

# make a list of building abbreviations and sort them into alphabetical order
# Hint: use the dictionary functions .values() or .keys() to list the values or keys, and then use the function 'sorted' to sort the list.
sorted_abv = ...

# make a dictionary that maps building names to the years occupied
building_year = ...

# make a dictionary that maps building names to zones
building_zone = ...

# use the dictionaries above to find the oldest building in the dictionary
old_building = ...
old_zone = ...
old_abv = ...

# use the dictionaries above to find the oldest building in zone 3
old_building2 = ...
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\n# find the abbreviations for the buildings\nbuildings = ["ATHLETIC FIELDS PAVILION (EASTSIDE)", \n            "ALMETRIS DUREN RESIDENCE HALL",\n            "ATHLETIC FIELDS PAVILION (REHAB)",\n            "ATHLETIC FIELDS PAVILION",\n            "ANNA HISS GYMNASIUM"]\n# make a dictionary that maps building names (key) to building abbreviations (value)\nbuilding_abv = ...\n\n# make a dictionary that maps building abbreviations (key) to building names (value)\n# Hint: use building_abv to create the this dictionary. Don\'t manually type keys and values...\nabv_building = ...\n\n# make a list of building abbreviations and sort them into alphabetical order\n# Hint: use the dictionary functions .values() or .keys() to list the values or keys, and then use the function \'sorted\' to sort the list.\nsorted_abv = ...\n\n# make a dictionary that maps building names to the years occupied\nbuilding_year = ...\n\n# make a dictionary that maps building names to zones\nbuilding_zon
```

```python
sorted_abv # should return ['ADH', 'AF1', 'AF2', 'AFP', 'AHG']
```

Output:
```
['ADH', 'AF1', 'AF2', 'AFP', 'AHG']
```

```python
old_building # should return 'ANNA HISS GYMNASIUM'
```

Output:
```
'ANNA HISS GYMNASIUM'
```

```python
old_abv # should return 'AHG'
```

Output:
```
'AHG'
```

```python
old_zone # should return 1
```

Output:
```
1
```

```python
old_building2 # should return 'ATHLETIC FIELDS PAVILION'
```

Output:
```
'ATHLETIC FIELDS PAVILION'
```

Tests for q1 (points: 2.5)
```
>>> a = [(1, 1), (2, 5), (3, 1)]
>>> area_triangle(a)
4.0
```
```
>>> a = [(1, 1), (2, 5)]
>>> area_triangle(a)
incorrect number of coordinates
```
```
>>> a = [(1, 1), (2, 5), (3, 1), (2, 0)]
>>> area_triangle(a)
incorrect number of coordinates
```
```
>>> a = [(1, 1), (2, 5), (2, 5)]
>>> area_triangle(a)
invalid coordinates for a triangle
```


Tests for q2 (points: 2.5)
```
>>> vertices = [(1, 1), (2, 5), (3, 1)]
>>> polygon_area(vertices)
4.0
```
```
>>> vertices = [(1, 1), (2, 5), (5, 5), (6, 1)]
>>> polygon_area(vertices)
16.0
```


Tests for q3 (points: 2.5)
```
>>> a = [0, 4.3, -5, 7]
>>> b = [-2, 7, -15, 1]
>>> sum_vector(a, b)
[-2, 11.3, -20, 8]
```


Tests for q4 (points: 2.5)
```
>>> sorted_abv
['ADH', 'AF1', 'AF2', 'AFP', 'AHG']
```
```
>>> old_building
'ANNA HISS GYMNASIUM'
```
```
>>> old_abv
'AHG'
```
```
>>> old_zone
1
```
```
>>> old_building2
'ATHLETIC FIELDS PAVILION'
```
