## Exercise 05 Error checking

### Exercise 05.0 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.


```python
# Your function should be able to "raise" error messages when the input is invalid.
# "raise" ValueError messages, not "print"
# "raise" ValueError messages, not "print"
# "raise" ValueError messages, not "print"
def fibonacci(n):
    ...
```

```python
# Perform some tests    
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(2) == 1
assert fibonacci(3) == 2
assert fibonacci(10) == 55
assert fibonacci(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    fibonacci(-1)
# Check that ValueError is raised for n is not an integer
with pytest.raises(ValueError):
    fibonacci(2.5)
```

## Exercise 05.1 (raising exceptions)

Write root finding functions using (bisection method, Newton Raphson and fixed point iteration) to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

Compute the roots of the following Polynomial function $ x^3 - 6x^2 + 4x + 12$  using:
* Bisection method
* Newton Raphson iterations


*HINT* <br>
raise ValueError for invalid initial guess(x0,x1) of bisection method<br>
raise RuntimeError for max iteration exceeded

```python
def fn(x):
    ...

# biscection
def bisection(f, x0, x1, tol, max_it):
    ...

# Newton-Raphson
def dfn(x):
    ...

def newton_raphson(f, df, x0, tol, max_it):
    ...

    
bisection(fn,4, 6,1e-6,100)
```

```python
# bisection
# Test with max_it = 30
x, error, num_it = bisection(fn, x0=3, x1=6, tol=1.0e-6, max_it=30)
assert round(x-4.5340702, 5) == 0
assert num_it > 20
```

```python
# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, error, num_it = bisection(fn, x0=3, x1=6, tol=1.0e-6, max_it=20)
```

```python
# Newton-Raphson
# Test with max_it = 30
x, error, num_it = newton_raphson(fn, dfn, x0=3, tol=1.0e-6, max_it=30)
assert round(x-2.51730405, 5) == 0
assert num_it > 4
```

```python
# Test with max_it to fail
with pytest.raises(RuntimeError):
    x, error, num_it = newton_raphson(fn, dfn, x0=3, tol=1.0e-6, max_it=3)
```

### Fixed-point iteration

Usually a formula for finding the root of an equation can be found by rearranging $f(x) = 0$ to be: $x = g(x)$ and then using the computation formula: $$ x_{i+1} = g(x_i)$$ to solve for succesively more accurate approximations of the root. 

Consider: $f(x) = x^2 - 4\sin(x) =0 $ can be rearranged as $x = g(x) = 4\frac{\sin(x)}{x}$. So the computational formula is $$ x_{i+1} = g(x_i) = 4\frac{\sin(x_i)}{x_i}$$

We could solve the equations for $x$ stating at an initial guess of $x = x_0$. The relative approximation error is computed as:

$\eta = \left|\frac{x_{i+1} - x_i}{x_{i+1}}\right| < \varepsilon$.

Compute one of the roots of function $ x^3 - 6x^2 + 4x + 12$ using Fixed-point iteration approach


```python
def fixed_point(f, x0, tol, max_it):
    ...
# find 1 of 3 roots with fixed point method
# you may try different x0 and max iterations to find the root
x, error, num_it = fixed_point(fn, x0=..., tol=1.0e-6, max_it=...)
root = ... # store the root in this variable
```

```python
import numpy as np
assert ((np.roots([1,-6,4,12])-root).round(4)==0).any()
```

```python
# Test with max_it to fail
with pytest.raises(RuntimeError):
    x, error, num_it = fixed_point(fn, x0=3, tol=1.0e-6, max_it=2)
```

Tests for q1 (points: 3.0)
```
>>> assert fibonacci(0) == 0
>>> assert fibonacci(1) == 1
>>> assert fibonacci(2) == 1
>>> assert fibonacci(3) == 2
>>> assert fibonacci(10) == 55
>>> assert fibonacci(15) == 610
>>> import pytest
>>> with pytest.raises(ValueError):
...     fibonacci(-1)
>>> with pytest.raises(ValueError):
...     fibonacci(2.5)
```


Tests for q2 (points: 4.0)
```
>>> (x, error, num_it) = bisection(fn, x0=3, x1=6, tol=1e-06, max_it=30)
>>> assert round(x - 4.5340702, 5) == 0
>>> assert num_it > 20
```
```
>>> with pytest.raises(RuntimeError):
...     (x, error, num_it) = bisection(fn, x0=3, x1=6, tol=1e-06, max_it=20)
```
```
>>> (x, error, num_it) = newton_raphson(fn, dfn, x0=3, tol=1e-06, max_it=30)
>>> assert round(x - 2.51730405, 5) == 0
>>> assert num_it > 4
```
```
>>> with pytest.raises(RuntimeError):
...     (x, error, num_it) = newton_raphson(fn, dfn, x0=3, tol=1e-06, max_it=3)
```


Tests for q3 (points: 4.0)
```
>>> import numpy as np
>>> assert ((np.roots([1, -6, 4, 12]) - root).round(4) == 0).any()
```
```
>>> with pytest.raises(RuntimeError):
...     (x, error, num_it) = fixed_point(fn, x0=3, tol=1e-06, max_it=2)
```
